<p align="right"><a href="./Gradient_Checking.ipynb">English</a> | <strong>中文</strong></p>

# 梯度检验（Gradient Checking）

欢迎来到本周的最后一个作业！在这个作业里，你将实现梯度检验（gradient checking）。

到本 notebook 结束时，你将能够：

实现梯度检验，来验证你反向传播实现的准确性。

## 关于提交给 AutoGrader 的重要提示

在把你的作业提交给 AutoGrader 之前，请确认你没有做以下这些事：

1. 你没有在作业里加任何 _额外的_ `print` 语句。
2. 你没有在作业里加任何 _额外的_ code cell。
3. 你没有更改任何函数参数。
4. 你没有在评分练习内部使用任何全局变量。除非题目特别要求，否则请避免这样做，改用局部变量。
5. 你没有在不需要改的地方更改作业代码，比如创建 _额外的_ 变量。

如果你做了上述任何一件事，提交作业时你会得到类似 `Grader Error: Grader feedback not found`（或其他类似意外）的错误。在寻求帮助/调试作业里的错误之前，请先检查这些。如果确实如此、而你又不记得自己改了什么，可以按这些[说明](https://www.coursera.org/learn/deep-neural-network/supplement/QWEnZ/h-ow-to-refresh-your-workspace)获取一份全新的作业副本。

## 目录
- [1 - 库 Packages](#1)
- [2 - 问题陈述](#2)
- [3 - 梯度检验是如何工作的？](#3)
- [4 - 一维梯度检验](#4)
    - [练习 Exercise 1 - forward_propagation](#ex-1)
    - [练习 Exercise 2 - backward_propagation](#ex-2)
    - [练习 Exercise 3 - gradient_check](#ex-3)
- [5 - N 维梯度检验](#5)
    - [练习 Exercise 4 - gradient_check_n](#ex-4)

<a name='1'></a>
## 1 - 库 Packages

```python
### v2.0
```

```python
import numpy as np
from testCases import *
from public_tests import *
from gc_utils import sigmoid, relu, dictionary_to_vector, vector_to_dictionary, gradients_to_vector

%load_ext autoreload
%autoreload 2
```

<a name='2'></a>
## 2 - 问题陈述

你是一个团队的一员，致力于让移动支付在全球可用，你被要求构建一个深度学习模型来检测欺诈——每当有人付款时，你想看看这笔付款是否可能是欺诈性的，比如用户的账户是否被黑客接管了。

你已经知道反向传播实现起来相当有挑战性，有时还有 bug。因为这是一个关乎成败的应用，你公司的 CEO 想真正确定你的反向传播实现是正确的。你的 CEO 说："给我证明你的反向传播确实在工作！" 为了给出这份保证，你将使用 "梯度检验"。

我们动手吧！

<a name='3'></a>
## 3 - 梯度检验是如何工作的？
反向传播计算梯度 $\frac{\partial J}{\partial \theta}$，其中 $\theta$ 表示模型的参数。$J$ 用前向传播和你的损失函数来计算。

因为前向传播相对容易实现，你有信心你把它做对了，所以你几乎 100% 确定你在正确地计算代价 $J$。因此，你可以用你计算 $J$ 的代码来验证计算 $\frac{\partial J}{\partial \theta}$ 的代码。

我们回过头看导数（或梯度）的定义：$$ \frac{\partial J}{\partial \theta} = \lim_{\varepsilon \to 0} \frac{J(\theta + \varepsilon) - J(\theta - \varepsilon)}{2 \varepsilon} \tag{1}$$

如果你不熟悉 "$\displaystyle \lim_{\varepsilon \to 0}$" 这个记号，它只是 "当 $\varepsilon$ 非常非常小时" 的一种说法。

你知道以下几点：

$\frac{\partial J}{\partial \theta}$ 是你想确保自己在正确计算的东西。
你能计算 $J(\theta + \varepsilon)$ 和 $J(\theta - \varepsilon)$（在 $\theta$ 是实数的情形），因为你确信你对 $J$ 的实现是正确的。
我们用方程 (1) 和一个小的 $\varepsilon$ 值，来让你的 CEO 相信你计算 $\frac{\partial J}{\partial \theta}$ 的代码是正确的！

<a name='4'></a>
## 4 - 一维梯度检验

考虑一个一维线性函数 $J(\theta) = \theta x$。这个模型只含单个实值参数 $\theta$，并以 $x$ 为输入。

你将实现代码来计算 $J(.)$ 及其导数 $\frac{\partial J}{\partial \theta}$。然后你会用梯度检验来确保你对 $J$ 的导数计算是正确的。

<img src="images/1Dgrad_kiank.png" style="width:600px;height:250px;">
<caption><center><font color='purple'><b>图 1</b>：一维线性模型 </font></center></caption>

上图展示了关键的计算步骤：先从 $x$ 开始，然后求值函数 $J(x)$（"前向传播"）。接着计算导数 $\frac{\partial J}{\partial \theta}$（"反向传播"）。

<a name='ex-1'></a>
### 练习 Exercise 1 - forward_propagation

实现 `forward propagation`。对这个简单函数，计算 $J(.)$

<!-- grade_id: cell-a4be88c5c0419ab7, points: 0, grade: false, solution: true -->
```python
# GRADED FUNCTION: forward_propagation

def forward_propagation(x, theta):
    """
    Implement the linear forward propagation (compute J) presented in Figure 1 (J(theta) = theta * x)
    
    Arguments:
    x -- a real-valued input
    theta -- our parameter, a real number as well
    
    Returns:
    J -- the value of function J, computed using the formula J(theta) = theta * x
    """
    
    # (approx. 1 line)
    # J = 
    # YOUR CODE STARTS HERE
    J = np.dot(theta,x)
    # YOUR CODE ENDS HERE
    
    return J
```

<!-- grade_id: cell-805a7fd19d554221, points: 10, grade: true, solution: false -->
```python
x, theta = 2, 4
J = forward_propagation(x, theta)
print ("J = " + str(J))
forward_propagation_test(forward_propagation)
```

Output:
```
J = 8
[92m All tests passed.
```

<a name='ex-2'></a>
### 练习 Exercise 2 - backward_propagation

现在，实现图 1 的 `backward propagation` 步骤（导数计算）。也就是，计算 $J(\theta) = \theta x$ 关于 $\theta$ 的导数。为了让你不用亲自算微积分，你应得到 $dtheta = \frac { \partial J }{ \partial \theta} = x$。

<!-- grade_id: cell-c06a1275399b210f, points: 0, grade: false, solution: true -->
```python
# GRADED FUNCTION: backward_propagation

def backward_propagation(x, theta):
    """
    Computes the derivative of J with respect to theta (see Figure 1).
    
    Arguments:
    x -- a real-valued input
    theta -- our parameter, a real number as well
    
    Returns:
    dtheta -- the gradient of the cost with respect to theta
    """
    
    # (approx. 1 line)
    # dtheta = 
    # YOUR CODE STARTS HERE
    dtheta = x
    # YOUR CODE ENDS HERE
    
    return dtheta
```

<!-- grade_id: cell-7b67ed84ac8bfd91, points: 10, grade: true, solution: false -->
```python
x, theta = 3, 4
dtheta = backward_propagation(x, theta)
print ("dtheta = " + str(dtheta))
backward_propagation_test(backward_propagation)
```

Output:
```
dtheta = 3
[92m All tests passed.
```

#### 期望输出：

```
dtheta = 3
 All tests passed.
```

<a name='ex-3'></a>
### 练习 Exercise 3 - gradient_check

为了表明 `backward_propagation()` 函数在正确地计算梯度 $\frac{\partial J}{\partial \theta}$，我们来实现梯度检验。

**说明**：
- 首先用上面的公式 (1) 和一个小的 $\varepsilon$ 值计算 "gradapprox"。以下是要遵循的步骤：
    1. $\theta^{+} = \theta + \varepsilon$
    2. $\theta^{-} = \theta - \varepsilon$
    3. $J^{+} = J(\theta^{+})$
    4. $J^{-} = J(\theta^{-})$
    5. $gradapprox = \frac{J^{+} - J^{-}}{2  \varepsilon}$
- 然后用反向传播计算梯度，把结果存进一个变量 "grad"
- 最后，用以下公式计算 "gradapprox" 和 "grad" 之间的相对差：
$$ difference = \frac {\mid\mid grad - gradapprox \mid\mid_2}{\mid\mid grad \mid\mid_2 + \mid\mid gradapprox \mid\mid_2} \tag{2}$$
你需要 3 个步骤来计算这个公式：
   - 1'. 用 np.linalg.norm(...) 计算分子
   - 2'. 计算分母。你需要调用 np.linalg.norm(...) 两次。
   - 3'. 把它们相除。
- 如果这个差很小（比如小于 $10^{-7}$），你就可以相当有信心你正确地计算了梯度。否则，梯度计算里可能有错误。

<!-- grade_id: cell-ed57ede577f9d607, points: 0, grade: false, solution: true -->
```python
# GRADED FUNCTION: gradient_check

def gradient_check(x, theta, epsilon=1e-7, print_msg=False):
    """
    Implement the gradient checking presented in Figure 1.
    
    Arguments:
    x -- a float input
    theta -- our parameter, a float as well
    epsilon -- tiny shift to the input to compute approximated gradient with formula(1)
    
    Returns:
    difference -- difference (2) between the approximated gradient and the backward propagation gradient. Float output
    """
    
    # Compute gradapprox using right side of formula (1). epsilon is small enough, you don't need to worry about the limit.
    # (approx. 5 lines)
    # theta_plus =                                 # Step 1
    # theta_minus =                                # Step 2
    # J_plus =                                    # Step 3
    # J_minus =                                   # Step 4
    # gradapprox =                                # Step 5
    # YOUR CODE STARTS HERE
    thetaplus = theta + epsilon                           # Step 1
    thetaminus = theta - epsilon                          # Step 2
    J_plus = np.dot(thetaplus,x)                            # Step 3
    J_minus = np.dot(thetaminus,x)                        # Step 4
    gradapprox = (J_plus - J_minus)/(2*epsilon)
    
    # YOUR CODE ENDS HERE
    
    # Check if gradapprox is close enough to the output of backward_propagation()
    #(approx. 1 line) DO NOT USE "grad = gradapprox"
    # grad =
    # YOUR CODE STARTS HERE
    grad = x
    # YOUR CODE ENDS HERE
    
    #(approx. 3 lines)
    # numerator =                                 # Step 1'
    # denominator =                               # Step 2'
    # difference =                                # Step 3'
    # YOUR CODE STARTS HERE
    numerator = np.linalg.norm(gradapprox-grad)                         # Step 1'
    denominator = np.linalg.norm(gradapprox) + np.linalg.norm(grad)     # Step 2'
    difference = numerator/denominator                                  # Step 3'
    # YOUR CODE ENDS HERE
    if print_msg:
        if difference > 2e-7:
            print ("\033[93m" + "There is a mistake in the backward propagation! difference = " + str(difference) + "\033[0m")
        else:
            print ("\033[92m" + "Your backward propagation works perfectly fine! difference = " + str(difference) + "\033[0m")
    
    return difference
```

<!-- grade_id: cell-be0338be7d50dd11, points: 10, grade: true, solution: false -->
```python
x, theta = 3, 4
difference = gradient_check(x, theta, print_msg=True)
gradient_check_test(gradient_check, difference)
```

Output:
```
[92mYour backward propagation works perfectly fine! difference = 7.814075313343006e-11[0m
[92m All tests passed.
```

**期望输出**：

<table>
    <tr>
        <td>  <b> Your backward propagation works perfectly fine!</b>  </td>
        <td> difference = 7.814075313343006e-11 </td>
    </tr>
</table>

恭喜，这个差小于 $2 * 10^{-7}$ 的阈值。所以你可以高度确信你在 `backward_propagation()` 里正确地计算了梯度。

现在，在更一般的情形，你的代价函数 $J$ 有不止一个一维输入。当你训练一个神经网络时，$\theta$ 实际上由多个矩阵 $W^{[l]}$ 和偏置 $b^{[l]}$ 组成！知道怎么对更高维的输入做梯度检验很重要。我们来做吧！

<a name='5'></a>
## 5 - N 维梯度检验

下图描述了你欺诈检测模型的前向和反向传播。

<img src="images/NDgrad_kiank.png" style="width:600px;height:400px;">
<caption><center><font color='purple'><b>图 2</b>：深度神经网络。LINEAR -> RELU -> LINEAR -> RELU -> LINEAR -> SIGMOID</font></center></caption>

我们来看看你对 `forward_propagation` 和 `backward_propagation` 的实现。

下面是为 `forward_propagation_n` 提供的代码。注意这里名字里的 `n` 表示它是给 n 维用的

```python
def forward_propagation_n(X, Y, parameters):
    """
    Implements the forward propagation (and computes the cost) presented in Figure 3.
    
    Arguments:
    X -- training set for m examples
    Y -- labels for m examples 
    parameters -- python dictionary containing your parameters "W1", "b1", "W2", "b2", "W3", "b3":
                    W1 -- weight matrix of shape (5, 4)
                    b1 -- bias vector of shape (5, 1)
                    W2 -- weight matrix of shape (3, 5)
                    b2 -- bias vector of shape (3, 1)
                    W3 -- weight matrix of shape (1, 3)
                    b3 -- bias vector of shape (1, 1)
    
    Returns:
    cost -- the cost function (logistic cost for m examples)
    cache -- a tuple with the intermediate values (Z1, A1, W1, b1, Z2, A2, W2, b2, Z3, A3, W3, b3)

    """
    
    # retrieve parameters
    m = X.shape[1]
    W1 = parameters["W1"]
    b1 = parameters["b1"]
    W2 = parameters["W2"]
    b2 = parameters["b2"]
    W3 = parameters["W3"]
    b3 = parameters["b3"]

    # LINEAR -> RELU -> LINEAR -> RELU -> LINEAR -> SIGMOID
    Z1 = np.dot(W1, X) + b1
    A1 = relu(Z1)
    Z2 = np.dot(W2, A1) + b2
    A2 = relu(Z2)
    Z3 = np.dot(W3, A2) + b3
    A3 = sigmoid(Z3)

    # Cost
    log_probs = np.multiply(-np.log(A3),Y) + np.multiply(-np.log(1 - A3), 1 - Y)
    cost = 1. / m * np.sum(log_probs)
    
    cache = (Z1, A1, W1, b1, Z2, A2, W2, b2, Z3, A3, W3, b3)
    
    return cost, cache
```

现在，我们来看反向传播的代码。

下面是为 `backward_propagation_n` 提供的代码。注意这里名字里的 `n` 表示它是给 n 维用的

```python
def backward_propagation_n(X, Y, cache):
    """
    Implement the backward propagation presented in figure 2.
    
    Arguments:
    X -- input datapoint, of shape (input size, 1)
    Y -- true "label"
    cache -- cache output from forward_propagation_n()
    
    Returns:
    gradients -- A dictionary with the gradients of the cost with respect to each parameter, activation and pre-activation variables.
    """
    
    m = X.shape[1]
    (Z1, A1, W1, b1, Z2, A2, W2, b2, Z3, A3, W3, b3) = cache
    
    dZ3 = A3 - Y
    dW3 = 1. / m * np.dot(dZ3, A2.T)
    db3 = 1. / m * np.sum(dZ3, axis=1, keepdims=True)
    
    dA2 = np.dot(W3.T, dZ3)
    dZ2 = np.multiply(dA2, np.int64(A2 > 0))
    dW2 = 1. / m * np.dot(dZ2, A1.T) * 2
    db2 = 1. / m * np.sum(dZ2, axis=1, keepdims=True)
    
    dA1 = np.dot(W2.T, dZ2)
    dZ1 = np.multiply(dA1, np.int64(A1 > 0))
    dW1 = 1. / m * np.dot(dZ1, X.T)
    db1 = 4. / m * np.sum(dZ1, axis=1, keepdims=True)
    
    gradients = {"dZ3": dZ3, "dW3": dW3, "db3": db3,
                 "dA2": dA2, "dZ2": dZ2, "dW2": dW2, "db2": db2,
                 "dA1": dA1, "dZ1": dZ1, "dW1": dW1, "db1": db1}
    
    return gradients
```

如果你刚实现了这些函数，你可能对它们是否正确工作没有高度信心。所以我们来实现梯度检验，帮忙验证其表现。

**梯度检验是如何工作的？**

和第 3、4 节一样，你想把 "gradapprox" 和反向传播算出的梯度做比较。公式仍是：

$$ \frac{\partial J}{\partial \theta} = \lim_{\varepsilon \to 0} \frac{J(\theta + \varepsilon) - J(\theta - \varepsilon)}{2 \varepsilon} \tag{1}$$

不过，$\theta$ 不再是一个标量了。它是一个叫 "parameters" 的字典。函数 "`dictionary_to_vector()`" 已为你实现好。它把 "parameters" 字典转换成一个叫 "values" 的向量，做法是把所有参数（W1, b1, W2, b2, W3, b3）都重塑成向量再拼接起来。

其逆函数是 "`vector_to_dictionary`"，它把 "parameters" 字典输出回来。

<img src="images/dictionary_to_vector.png" style="width:600px;height:400px;">
<caption><center><font color='purple'><b>图 2</b>：dictionary_to_vector() 和 vector_to_dictionary()。你会在 gradient_check_n() 里需要这些函数</font></center></caption>

"gradients" 字典也已用 gradients_to_vector() 转换成了一个向量 "grad"，所以你不用为那个操心。

现在，对你向量里的每一个参数，你都会施加和 gradient_check 练习一样的过程。你会把每个梯度近似存进一个向量 `gradapprox`。如果检验如期进行，这个近似里的每个值都必须与存在 `grad` 向量里的真实梯度值匹配。

注意 `grad` 是用函数 `gradients_to_vector` 计算的，它用了 `backward_propagation_n` 函数的梯度输出。

<a name='ex-4'></a>
### 练习 Exercise 4 - gradient_check_n

实现下面的函数。

**说明**：这里是能帮你实现梯度检验的伪代码。

对 num_parameters 里的每个 i：
- 要计算 `J_plus[i]`：
    1. 把 $\theta^{+}$ 设为 `np.copy(parameters_values)`
    2. 把 $\theta^{+}_i$ 设为 $\theta^{+}_i + \varepsilon$
    3. 用 `forward_propagation_n(x, y, vector_to_dictionary(`$\theta^{+}$ `))` 计算 $J^{+}_i$。
- 要计算 `J_minus[i]`：用 $\theta^{-}$ 做同样的事
- 计算 $gradapprox[i] = \frac{J^{+}_i - J^{-}_i}{2 \varepsilon}$

这样，你得到一个向量 gradapprox，其中 gradapprox[i] 是关于 `parameter_values[i]` 的梯度的近似。你现在可以把这个 gradapprox 向量和反向传播来的 gradients 向量做比较。就像一维情形（步骤 1'、2'、3'）那样，计算：
$$ difference = \frac {\| grad - gradapprox \|_2}{\| grad \|_2 + \| gradapprox \|_2 } \tag{3}$$

**注意**：用 `np.linalg.norm` 来求范数

<!-- grade_id: cell-1e5a768bc4e28e66, points: 0, grade: false, solution: true -->
```python
# GRADED FUNCTION: gradient_check_n

def gradient_check_n(parameters, gradients, X, Y, epsilon=1e-7, print_msg=False):
    """
    Checks if backward_propagation_n computes correctly the gradient of the cost output by forward_propagation_n
    
    Arguments:
    parameters -- python dictionary containing your parameters "W1", "b1", "W2", "b2", "W3", "b3"
    grad -- output of backward_propagation_n, contains gradients of the cost with respect to the parameters 
    X -- input datapoint, of shape (input size, number of examples)
    Y -- true "label"
    epsilon -- tiny shift to the input to compute approximated gradient with formula(1)
    
    Returns:
    difference -- difference (2) between the approximated gradient and the backward propagation gradient
    """
    
    # Set-up variables
    parameters_values, _ = dictionary_to_vector(parameters)
    
    grad = gradients_to_vector(gradients)
    num_parameters = parameters_values.shape[0]
    J_plus = np.zeros((num_parameters, 1))
    J_minus = np.zeros((num_parameters, 1))
    gradapprox = np.zeros((num_parameters, 1))
    
    # Compute gradapprox
    for i in range(num_parameters):
        
        # Compute J_plus[i]. Inputs: "parameters_values, epsilon". Output = "J_plus[i]".
        # "_" is used because the function you have outputs two parameters but we only care about the first one
        #(approx. 3 lines)
        # theta_plus =                                        # Step 1
        # theta_plus[i] =                                     # Step 2
        # J_plus[i], _ =                                     # Step 3
        # YOUR CODE STARTS HERE
        thetaplus = np.copy(parameters_values)                                       # Step 1
        thetaplus[i][0] = thetaplus[i][0] + epsilon                                  # Step 2
        J_plus[i], _ = forward_propagation_n(X, Y, vector_to_dictionary(thetaplus))  # Step 3
        # YOUR CODE ENDS HERE
        
        # Compute J_minus[i]. Inputs: "parameters_values, epsilon". Output = "J_minus[i]".
        #(approx. 3 lines)
        # theta_minus =                                    # Step 1
        # theta_minus[i] =                                 # Step 2        
        # J_minus[i], _ =                                 # Step 3
        # YOUR CODE STARTS HERE
        thetaminus = np.copy(parameters_values)                                       # Step 1
        thetaminus[i][0] = thetaminus[i][0] - epsilon                                 # Step 2
        J_minus[i], _ = forward_propagation_n(X, Y, vector_to_dictionary(thetaminus)) # Step 3
        # YOUR CODE ENDS HERE
        
        # Compute gradapprox[i]
        # (approx. 1 line)
        # gradapprox[i] = 
        # YOUR CODE STARTS HERE
        gradapprox[i] = (J_plus[i] - J_minus[i])/(2*epsilon)
        # YOUR CODE ENDS HERE
    
    # Compare gradapprox to backward propagation gradients by computing difference.
    # (approx. 3 line)
    # numerator =                                             # Step 1'
    # denominator =                                           # Step 2'
    # difference =                                            # Step 3'
    # YOUR CODE STARTS HERE
    numerator = np.linalg.norm(grad - gradapprox)                       # Step 1'
    denominator = np.linalg.norm(grad) + np.linalg.norm(gradapprox)     # Step 2'
    difference = numerator/denominator                                  # Step 3'
    # YOUR CODE ENDS HERE
    if print_msg:
        if difference > 2e-7:
            print ("\033[93m" + "There is a mistake in the backward propagation! difference = " + str(difference) + "\033[0m")
        else:
            print ("\033[92m" + "Your backward propagation works perfectly fine! difference = " + str(difference) + "\033[0m")

    return difference
```

<!-- grade_id: cell-0d7896ce7c954fc9, points: 20, grade: true, solution: false -->
```python
X, Y, parameters = gradient_check_n_test_case()

cost, cache = forward_propagation_n(X, Y, parameters)
gradients = backward_propagation_n(X, Y, cache)
difference = gradient_check_n(parameters, gradients, X, Y, 1e-7, True)
expected_values = [0.2850931567761623, 1.1890913024229996e-07]
assert not(type(difference) == np.ndarray), "You are not using np.linalg.norm for numerator or denominator"
assert np.any(np.isclose(difference, expected_values)), "Wrong value. It is not one of the expected values"

```

Output:
```
[93mThere is a mistake in the backward propagation! difference = 0.2850931567761623[0m
```

**期望输出**：

<table>
    <tr>
        <td>  <b> There is a mistake in the backward propagation!</b>  </td>
        <td> difference = 0.2850931567761623 </td>
    </tr>
</table>

看起来 `backward_propagation_n` 代码里有错误！幸好你实现了梯度检验。回到 `backward_propagation_n`，试着找出/改正这些错误 *（提示：检查 dW2 和 db1）*。当你认为改好了，就重跑梯度检验。记住，如果你改了代码，你需要重新执行定义 `backward_propagation_n()` 的那个 cell。

你能让梯度检验宣告你的导数计算正确吗？尽管这部分作业不评分，你也应该试着找出 bug 并重跑梯度检验，直到你确信反向传播现在被正确地实现了。

**注意**
- 梯度检验很慢！用 $\frac{\partial J}{\partial \theta} \approx  \frac{J(\theta + \varepsilon) - J(\theta - \varepsilon)}{2 \varepsilon}$ 来近似梯度在计算上代价高昂。因此，我们不会在训练时的每次迭代都跑梯度检验。只跑几次，检查梯度是否正确即可。
- 梯度检验，至少按我们呈现的方式，对 dropout 不起作用。你通常会在不带 dropout 的情况下跑梯度检验算法来确保你的反向传播正确，然后再加上 dropout。

恭喜！现在你可以确信你用于欺诈检测的深度学习模型在正确工作了！你甚至可以用这个来说服你的 CEO。:)
<br>
<font color='blue'>

**你应该从这个 notebook 记住的**：
- 梯度检验验证反向传播来的梯度与梯度的数值近似（用前向传播计算）之间的接近程度。
- 梯度检验很慢，所以你不会想在训练的每次迭代都跑它。你通常只会跑它来确保你的代码正确，然后关掉它、用反向传播来做实际的学习过程。

```python

```